# <div style="background-color:rgba(204, 229, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(0, 76, 153, 1);">PHYS 231 Lab #0<span style="color:red"> $\to$ (10 possible marks)</span> </span></div> 

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-231-images/main/Lab-0/images/PHYS231.png" alt="image info" style="width: 85%; height: auto;" />
</center>
</p>

<font size ="4"><span style="color:purple">We recommend that you complete this notebook using either the [Google Chrome](https://www.google.com/intl/en_ca/chrome/) or [Mozilla Firefox](https://www.mozilla.org/en-CA/firefox/new/) browser.  Chrome and Firefox are both available for Windows, Mac, and Linux operating systems.  The PHYS 121 Jupyter notebooks should work in other browsers, but we have specifically verified that they work well in both Chrome and Firefox.</span></font>

***
## Lab Learning Outcomes:
By the end of the PHYS 231 lab, the successful student will be able to:
* <b><span style="color:rgba(0, 153, 76, 1);">Design, build, and test practical analog circuits using passive components such as resistors, capacitors, inductors, and diodes.</span></b>
* <b><span style="color:rgba(0, 153, 76, 1);">Design, build, and test practical analog circuits using active components such as transistors and op amps.</span></b>
* <b><span style="color:rgba(0, 153, 76, 1);">Design, build, and test practical digital circuits using logic gates.</span></b>
* <b><span style="color:rgba(0, 153, 76, 1);">Set up an oscilloscope and use it to debug and analyze analog and digital circuits.</span></b>
* <b><span style="color:rgba(0, 153, 76, 1);">Use Python and Jupyter Notebooks to analyze data collected from the circuits that you design and build.</span></b>

***
## Lab Logistics:

The lab schedule and lab manuals are available on the course website: https://cmps-people.ok.ubc.ca/jbobowsk/phys231calendar.html

You will need to purchase **two** bound notebooks for use in the PHYS 231 labs.  Suitable lab notebooks that are available from the UBC Okanagan bookstore, from cheapest to most expensive, are shown in Fig. 1. 

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-231-images/main/Lab-0/images/yellow.jpg" alt="image info" style="width: 25%; height: auto;" /> &nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp; <img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-231-images/main/Lab-0/images/blue.jpg" alt="image info" style="width: 25%; height: auto;" /> &nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp; <img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-231-images/main/Lab-0/images/computation.jpg" alt="image info" style="width: 25.3%; height: auto;" /><br>
<b>Fig. 1: Suitable bound lab notebooks available in the UBC Okanagan bookstore.</b></center>
</p>

Please do not use coil bound notebooks, notepads, or the ubiquitous Hilroy exercise books as PHYS 231 lab notebooks (see Fig. 2).

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-231-images/main/Lab-0/images/coil.jpg" alt="image info" style="width: 25%; height: auto;" /> &nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp; <img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-231-images/main/Lab-0/images/notepad.jpg" alt="image info" style="width: 22.7%; height: auto;" />&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp; 
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-231-images/main/Lab-0/images/hilroy.jpg" alt="image info" style="width: 24%; height: auto;" /><br>
<b>Fig. 2: Notebooks that are not suitable for use in the PHYS 231 lab.</b></center>
</p>

In the PHYS 231 lab, you will:
- Submit the lab notebook (say, notebook No. 1) used to complete Lab $N-1$ at the *start* of Lab $N$.
- Complete Lab $N$ using your other notebook (notebook No. 2).
- Keep notebook No. 2 with you and complete your write up before the start of Lab $N+1$.
- Submit notebook No. 2 at the *start* of Lab $N+1$.  The graded notebook No. 1 will be returned to you and you will use it to complete Lab $N+1$. 

***
## Restart the Kernel:

To avoid unnecessary errors, **please begin by restarting the "kernel"**.  Restart the kernal by selecting **Kernel** $\to$ **Restart Kernel** from the menu options at to the top of the page.  Here is a <a href = "https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/gifs/restart.gif">GIF</a> showing how to complete these steps.  You don't need to execute any cells before the restart.  Simply go directly to the **Kernel** dropdown menu.

***
## Autograding:
The PHYS 231 Jupyter Notebooks will make use of some autograding.  To make the autograding work, the two cells below need to be executed.  The first one installs any missing required packages and the second imports required packages/modules.  If 'PHYS231.Installer()' reports that some functions have been installed, the user should execute the PHYS231.Installer() cell a second time.  The second time the installer function is run, it should report that **"All packages already installed. Please proceed"**.

The 'PHYS231.Installer()' command requires the file 'PHYS231.py', which you should see included in the list of files along the left-hand side of the screen.

```python
# Hide unnecessary warnings.
import warnings
warnings.filterwarnings('ignore')

# Import PHYS231.py and then run the installer function.
import PHYS231
PHYS231.Installer()
```

```python
# Initialize Otter
import otter
grader = otter.Notebook("PHYS 231 - Lab 0.ipynb")
```

***
## Import Modules:
Execute the cell below to import a number of useful pre-built Python modules that will be used in the PHYS 231 labs.

```python
import numpy as np
import matplotlib.pyplot as plt
import pandas as pd
from matplotlib.pyplot import cm # used to generate a sequence of colours for plotting
from scipy.optimize import curve_fit
from IPython.display import HTML as html_print
from IPython.display import display, Markdown, Latex
from IPython.display import YouTubeVideo
import math
import sys
```

```python
PHYS231.graderCheck(['PHYS231' in sys.modules, 'PHYS231' in sys.modules, 'numpy' in sys.modules], ['PHYS231', 'otter', 'numpy'], grader.check('q0.1'))
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Name & Student Number</span></div>  

So that we know who you were working with, both your name and student number and your partner's name and student number need to appear in the notebooks used during the labs.  In the cells below this text, enter your name and student number and your partner's name and student number as strings (i.e. place your names and student numbers between single or double quotes.

Examples of valid entries are:
```python
last_name_1 = 'Doe'
first_name_1 = 'Jane'
student_number_1 = '12345678'
```
and
```python
last_name_1 = "Doe"
first_name_1 = "John"
student_number_1 = "12345678"
```

<!-- BEGIN QUESTION -->

***
### **<span style="color:blue">Partner #1</span>** **<span style="color:red">(0.5 marks)</span>** 

```python
# Enter LAST name of partner 1.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
last_name_1 = ...
```

```python
# Enter FIRST name of partner 1.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
first_name_1 = ...
```

```python
# Enter STUDENT NUMBER of partner 1.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
student_number_1 = ...
```

***
### **<span style="color:blue">Partner #2</span>** **<span style="color:red">(0.5 marks)</span>** 
If you are working on your own, leave all of last_name_2, first_name_2, and student_number_2 as ...  Please still exectute these cells and then run the 'grader.check("name_and_student_number_2")' cell.

```python
# Enter LAST name of partner 2.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
last_name_2 = ...
```

```python
# Enter FIRST name of partner 2.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
first_name_2 = ...
```

```python
# Enter STUDENT NUMBER of partner 2.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
student_number_2 = ...
```

```python
PHYS231.graderCheck([last_name_1, first_name_1, student_number_1, last_name_2, first_name_2, student_number_2], ['last_name_1', 'first_name_1', 'student_number_1', 'last_name_2', 'first_name_2', 'student_number_2'], grader.check('name_and_student_number'))
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 1 - Using Functions from PHYS231.py</span></div>

<!-- BEGIN QUESTION -->

## Pandas Dataframe

Pandas is a powerful Python module that can be used to handle and manipulate small and large datasets.  In this introductory lab, we will use Pandas to import data stored in a "csv" ('**c**omma-**s**eparated **v**alues') file called "Lab0_data.csv".  The syntax for entering loading the data in the file is as follows:

```python
filename = 'name_of_file.csv'
df = pd.read_csv(filename)

# Display the dataframe (df)
df
```

***
**<span style="color:blue">Question 1.1</span>** **<span style="color:red">(1 mark)</span>**  

In the cell below, replace the ... with the name of the csv file between single or double quotes and then execute the cell.

*** Please do not change any of the variable names in the cell below ***

```python
filename = ...
df = pd.read_csv(filename)

# Display the dataframe
df
```

```python
PHYS231.graderCheck([filename], ['filename'], grader.check('q1.1'))
```

Once a dataframe has been created, you can access individual columns using the following notation:

```python
df['column_label']
```

where ```column_label``` should be replaced with the name of the column you're trying to access. If desired, you can also convert the entries in a dataframe column to an array via:

```python
array = df['column_label'].to_numpy()
```

***
**<span style="color:blue">Question 1.2</span>** **<span style="color:red">(3 marks)</span>**  

In the cell below, replace the ... with the three column names of the dataframe you created above and then execute the cell.

*** Please do not change any of the variable names in the cell below ***

```python

# time array
t = df['...'].to_numpy()

# y array
y = df['...'].to_numpy()

# erry array
erry = df['...'].to_numpy()

# Display the arrays
print('t =', t, 's')
print('y =', y, 'm')
print('erry =', erry, 'm')
```

```python
PHYS231.graderCheck([t, y, erry], ['t', 'y', 'erry'], grader.check('q1.2'))
```

<!-- BEGIN QUESTION -->

## Scatter Plots

We can create a scatter plot of the data loaded into out dataframne or arrays.  'PHYS231.py' contains a function called 'Scatter()'.  The most basic usage of this function is:
```python
fig1 = PHYS231.Scatter(t, y)
```
which will produce a scatter plot of column y versus column t.  If we add a third argument in the function call, it will produce a scatter plot with error bars.
```python
fig2 = PHYS231.Scatter(t, y, erry)
```
Additional options can be included to do some basic formatting.  Executing the following line of code:
```python
fig3 = PHYS231.Scatter(xData = t, yData = y, yErrors = erry, xlabel = 'time', ylabel = 'position', xUnits = 's', yUnits = 'm')
```
will produce a scatter plot with error bars and labeled $x$- and $y$-axes.

Each of these examples is demonstrated below.  

For additional information about the 'Scatter' function, see **Appendix A** near the end of this notebook. 

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 1.3</span>** **<span style="color:red">(2 marks)</span>**  

Replace all instances of ... in the cell below to produce:

 - (a) a plot of y versus t without error bars
 - (b) a plot of y versus t with error bars
 
 *** Please do not change any of the variable names in the cell below ***

```python

# (a) y versus t without error bars
fig1 = PHYS231.Scatter(..., ...)

# (a) y versus t with error bars
fig2 = PHYS231.Scatter(..., ..., ...)
```

<!-- END QUESTION -->

The cell below shows how custum axis labels can be added to the plots using the PHYS231.Scatter() function.  Execute the cell below to see the resulting output.

```python
fig3 = PHYS231.Scatter(
    xData = t,
    yData = y,
    yErrors = erry,
    xlabel = "time",
    ylabel = "position",
    xUnits = "s",
    yUnits = "m",
)
```

<!-- END QUESTION -->

## Linear Fits

The final thing that we'll do in Part 1 is fit a straight line through the data shown in the scatter plots above.  Of course, we don't need to use a pre-built function from PHYS231.py to do the fit.  If you're interested in seeing how to write your own code to do linear fits, see the following tutorial: https://cmps-people.ok.ubc.ca/jbobowsk/Python/html/Jupyter%20Weighted%20Linear%20Fit.html. 

For some common data analysis tasks, such as linear fits, it can be convenient to create functions to complete these repetitive tasks.  The PHYS 231 function that we'll use is called 'LinearFit'.  The most basic call of this function is:
```python
m, b, dm, db, fig = PHYS231.LinearFit(t, y)
```
This results in an *unweighted* fit to the data.  In an unweighted fit, the error bars are ignored when determining the slope and intercept of the best-fit line.

In an weighted fit, points with smaller error bars play a larger role in determining the slope and intercept of the best-fit line.  The function call to perform a weighted fit is:
```python
m, b, dm, db, fig = PHYS231.LinearFit(t, y, erry)
```
A weighted fit with a formatted plot can be called by using some of the other available arguments:
```python
m, b, dm, db, fig = PHYS231.LinearFit(xData = t, yData = y, yErrors = erry, xlabel = 'time', ylabel = 'position', xUnits = 's', yUnits = 'm')
```
For additional information about the 'LinearFit' function, see **Appendix B** near the end of this notebook. 

***
**<span style="color:blue">Question 1.4</span>** **<span style="color:red">(2 marks)</span>**  

Replace all instances of ... in the cell below to produce a weighted linear fit to the $y\pm\Delta y$ versus $t$ data.

 *** Please do not change any of the variable names in the cell below ***

```python

m, b, dm, db, fig4 = PHYS231.LinearFit(
    xData = ...,
    yData = ...,
    yErrors = ...,
    xlabel="time",
    ylabel="position",
    xUnits="s",
    yUnits="m",
)

```

```python
PHYS231.graderCheck([m, dm, b, db], ['m', 'dm', 'b', 'db'], grader.check('q1.4'))
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 2 - Basic Operations in Python</span></div>

For a set of Python tutorials that are primarily focused on data analysis, see: https://cmps-people.ok.ubc.ca/jbobowsk/Python.html

There is a large community of Python users.  Everything I learned about Python, and Jupyter Notebooks, I picked up from Stack Exchange and Stack Overflow.

Some things in Python can be done in exactly the way you would guess.  For example, the cells below demonstrate addition, subtraction, multiplication, and division. 

```python
3 + 8 - 2
```

```python
3 * 8 / 2
```

Other operations require that we make use of modules.  One of common Python module is 'Numpy'.  It is usually import as:
```python
import numpy as np
```
This allows use to call functions from the module using 'np' rather than typing all of 'numpy'.  Note that in this notebook, this import statement was already executed in the 'Import Modules' section and so does not need to be run again.  Numpy allows us to use trig functions, exponentials and logs, constants such as pi, and many other things.  If you were to try entering something like:
```python
cos(pi)
```
into a code cell, it would result in an error.  Feel free to try it!  This operation, however, can be done using numpy as demonstrated in the code cell below.

```python
np.cos(np.pi)
```

Here are some other examples.  Note that, in Python, ```np.log(x)``` is equivalent to $\ln(x)$ and ```np.log10(y)``` is equivalent to what we would usually write as $\log(y)$.

```python
np.exp(1)
```

```python
np.log(np.exp(1))
```

```python
np.log10(100)
```

```python
np.arctan(1) / np.pi
```

We'll discuss only one more example of standard mathematical operations.  In Python, to raise a quantity to a power, you can use the double multiplication sign notation as in:
```python
x**y
```
which is equivalent to $x^y$.

***
**<span style="color:blue">Question 2.1</span>** **<span style="color:red">(1 mark)</span>**  

Replace ... in the cell below with an expression that will calculate ans = $2^3$.

 *** Please do not change any of the variable names in the cell below ***

```python
ans = ...
print(ans)
```

```python
PHYS231.graderCheck([ans], ['ans'], grader.check('q2.1'))
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 3 - Feedback and Submission</span></div>

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 3.1:</span>**  

We welcome your feedback on the PHYS 231 labs!  Please feel free to include any comments you have about this lab in the cell below.  Your comments will be taken into consideration when revising/improving the PHYS 231 labs.  You can suggest improvements, point out anything that was unclear, comment on the strengths and weaknesses of the lab, ...

This question is optional and will have no impact on your lab grade.

***
**<span style="color:blue">Answer 3.1:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 


<!-- END QUESTION -->

***
Once you've completed this notebook:
- Save your work.
- Run 'grader.check_all()' to confirm that you've completed all required tasks.
- Run 'grader.export()' to generate a .zip file containing all of the materials that you will submit.
- Download the generated .zip file using the link that appears **below** the ```grader.check_all(...)``` cell.
- $\Leftarrow$ Please do **NOT** download the zip file from the frame along the left-hand side of the screen. 
- Upload the .zip file to the PHYS 121 Canvas gradebook.
- **Do NOT change the name of the .zip file.**
- **Do NOT modify the contents of the .zip file.**

Here is a <a href = "https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/gifs/Submission.gif">GIF</a> showing how these steps are completed.  Once your completed notebook has been uploaded to the Canvas gradebook, you're done!

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
display(grader.check_all())
PHYS231.saveMessage()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
PHYS231.waitMessage()
grader.export(files=["PHYS121_DataLogger.txt"])
PHYS231.proceedMessage()
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 4 - Playground (optional)</span></div>

Feel free to add as many cells as you like below and use them as a playground for further independent investigations.  These cells won't be graded, so feel free to use them in any way that you like.  

```python
# Here's an empty code cell that you can use.
```

```python
# Here's another empty code cell that you can use.
```

```python
# Here's yet another empty code cell that you can use.
# If you need more, you can add cells using the '+' icon in the menu bar at to the top of the screen.
```

### <div style="background-color:rgba(255, 255, 102, 0.5); text-align:left; padding:20px 0; margin-top:20px">$\quad$Appendix A &ndash; The Scatter Function...</div>

The function for generating scatter plots is called as follows:
```python
PHYS231.Scatter(xData, yData, yErrors = [], xlabel = 'x-axis', ylabel = 'y-axis', xUnits = '', yUnits = '', fill = False, show = True)
```
The 'xData' and 'yData' inputs are required, all other arguments are optional with default values set.  The function returns the a single output (the formatted plot):
```python
fig
```

The function will do a simple scatter plot if no 'yError' are included.  It will included error bars if 'yErrors' are passed to the function.  The 'fill' and 'show' arguments should generally be false.  

If the 'xData' list is empty, as in:
```python
xData = []
```
then the x-axis will be the trial number.  For example, is 'xData' is empty and 'yData' has 10 elements, the x-axis will span 1 to 10.

### Scatter Example Implmentation
The code block below shows an implementation of 'Scatter'.
```python
import PHYS231
theta = [10, 20, 30] # degrees
T = [2.02, 1.95, 2.13] # s
errT = [0.02, 0.02, 0.03] # s
fig = PHYS231.Scatter(theta, T, errT, 'initial angle' , 'period', 'degrees', 's')
```

If you're interested in generating your own plots with customized formatting, see the following Python-based plotting tutorial: https://cmps-people.ok.ubc.ca/jbobowsk/Python/html/Jupyter%20Basic%20Plots.html.
***

### <div style="background-color:rgba(255, 255, 102, 0.5); text-align:left; padding:20px 0; margin-top:20px">$\quad$Appendix B &ndash; The LinearFit Function...</div>

The linear-fit function is called as follows:
```python
PHYS231.LinearFit(xData, yData, yErrors = [], xlabel = 'x-axis', ylabel = 'y-axis', xUnits = '', yUnits = '', fill = False)
```
The xData and yData lists are required.  The others are optional with default values set.  The function returns the following outputs:
```python
Slope, Yintercept, errSlope, errYintercept, fig
```
The first four outputs are numerical values and 'fig' is the formatted plot. 

If $y$-uncertainities are provided, the function will perform a weighted fit.  The $y$-uncertainties list must be the same length as $x$- and $y$-data lists.  If $y$-uncertainties are not provided, the fit will be unweighted.  The other optional arguments include $x$- and $y$-axis names and units for the $x$- and $y$-datasets.  These must be entered as strings (enclosed in quotations) and they are used for formatting the outputs of the function.  If 'fill' is set to True, shading will be added around the best-fit line representing confindence intervals. 

### Unweighted Fit Example Implmentation
The code block below shows the most basic use of 'LinearFit' for an unweighted fit without any of the optional arguments.
```python
import PHYS231
V = [1, 2, 3, 4]
I = [0.12, 0.198, 0.285, 0.412]
m, b, dm, db, fig = PHYS231.LinearFit(V, I)
```

### Weighted Fit Example Implmentation
This second blcok of code shows how to use 'LinearFit' to do a weighted fit.  It also makes use of all of the other optional arguments.  
```python
import PHYS231
V = [1, 2, 3, 4]
I = [0.12, 0.198, 0.285, 0.412]
errI = [0.005, 0.012, 0.020, 0.025]
m, b, dm, db, fig = PHYS231.LinearFit(V, I, errI, 'voltage' , 'current', 'V', 'A', True)
```

The 'm, b, dm, db, fig = ...' syntax is used to store the function outputs in varibles.  If you're interested in details of the fitting process, see the following Python-based tutorial: https://cmps-people.ok.ubc.ca/jbobowsk/Python/html/Jupyter%20Weighted%20Linear%20Fit.html
***

<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/images/ubc-logo-full.jpg" width="500"/>

Last update: September 14, 2024

Tests for name_and_student_number (points: None)
```
>>> isinstance(last_name_1, str)
True
```
failure_message: last_name_1 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(first_name_1, str)
True
```
failure_message: first_name_1 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(student_number_1, str)
True
```
failure_message: student_number_1 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> test = True
>>> if isinstance(student_number_1, str):
...     test = len(student_number_1) == 8
>>> test
True
```
failure_message: student_number_1 must be a string of exactly 8 characters.
```
>>> test = True
>>> if isinstance(student_number_1, str):
...     test = student_number_1.isdigit()
>>> test
True
```
failure_message: student_number_1 must contain only interger digits (0, 1, 2, 3, ..., 9)
```
>>> isinstance(last_name_2, str) or last_name_2 == ...
True
```
failure_message: last_name_2 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(first_name_2, str) or first_name_2 == ...
True
```
failure_message: first_name_2 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(student_number_2, str) or student_number_2 == ...
True
```
failure_message: student_number_2 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> test = True
>>> if isinstance(student_number_2, str):
...     test = len(student_number_2) == 8
>>> test
True
```
failure_message: student_number_2 must be a string of exactly 8 characters or be left as ...
```
>>> test = True
>>> if isinstance(student_number_2, str):
...     test = student_number_2.isdigit()
>>> test
True
```
failure_message: student_number_2 must contain only interger digits (0, 1, 2, 3, ..., 9) or be left as ...


Tests for q0.1 (points: None)
```
>>> 'PHYS231' in sys.modules and 'otter' in sys.modules and ('numpy' in sys.modules)
True
```
failure_message: Make sure you execute the cells containing "import PHYS231" and "import otter" as well as the cell beneath the heading "Import Modules".


Tests for q1.1 (points: None)
```
>>> isinstance(filename, str)
True
```
failure_message: Your answer must be a string, i.e. surrounded by single or double quotes.
```
>>> filename == 'Lab0_data.csv'
True
```
failure_message: Please check the filename that you entered.  It must exactly match 'Lab0_data.csv' or "Lab0_data.csv".


Tests for q1.2 (points: None)
```
>>> all(np.isclose(t, df['t'], rtol=0.001))
True
```
failure_message: Make sure 't' array matches the 't' column in the dataframe.
```
>>> all(np.isclose(y, df['y'], rtol=0.001))
True
```
failure_message: Make sure 'y' array matches the 'y' column in the dataframe.
```
>>> all(np.isclose(erry, df['erry'], rtol=0.001))
True
```
failure_message: Make sure 'erry' array matches the 'erry' column in the dataframe.


Tests for q1.4 (points: None)
```
>>> math.isclose(m, 1.974, rel_tol=0.01)
True
```
failure_message: Your slope is incorrect.
```
>>> math.isclose(dm, 0.059, rel_tol=0.01)
True
```
failure_message: The error in your slope is incorrect.
```
>>> math.isclose(b, 4.23, rel_tol=0.01)
True
```
failure_message: Your intercept is incorrect.
```
>>> math.isclose(db, 0.1515, rel_tol=0.01)
True
```
failure_message: The error in your intercept is incorrect.


Tests for q2.1 (points: None)
```
>>> math.isclose(ans, 2 ** 3, rel_tol=1e-05)
True
```
failure_message: Please try again.
